```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment6.ipynb")
```

# k-Nearest Neighbors, Decision Trees, and Random Forests

In this assignment, you'll explore Decision Trees and Random Forests using the sklearn library. Follow the instructions, run all the cells, and submit the completed notebook via Gradescope.

For multiple choice questions, make sure to only include your final answer. For all code set `random_state=42`.

Throughout the assignment, you'll learn how to:

- Build a Decision Tree
- Compute the Gini index
- Construct a Random Forest and adjust its parameters to optimize performance

Let's begin by loading the necessary libraries and dataset.


```python
# Import necessary libraries
import numpy as np
import pandas as pd
from sklearn.model_selection import train_test_split
from sklearn.tree import DecisionTreeClassifier, export_text
from sklearn.ensemble import RandomForestClassifier
from sklearn.metrics import accuracy_score

# Load a dataset (Iris dataset)
from sklearn.datasets import load_iris
data = load_iris()
X = data.data
y = data.target

# Split the dataset into training and testing sets
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.5, random_state=42)

# Show dataset features and labels
print("Features:", data.feature_names)
print("Labels:", data.target_names)

```

<!-- BEGIN QUESTION -->

#### 1. What is the primary difference between classification and regression problems?

        A) Classification predicts continuous values, and regression predicts categories.
        B) Classification predicts categories, and regression predicts continuous values.
        C) Both are used to predict categories.
        D) Both are used to predict numeric values.

```python
def ans1():
    ...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### What is the purpose of the "splitting criterion" in a decision tree?

        A) To evaluate when to stop splitting
        B) To determine the optimal depth of the tree
        C) To select the best feature to split on
        D) To calculate the number of leaves

```python
def ans2():
    ...
```

<!-- END QUESTION -->


### 3. Building a Decision Tree


#### 3. Building a Decision Tree:
Create a Decision Tree and make sure the accuracy is above 90%. Also print out the tree with export_text(). Make sure to set the `random_state`.

```python
from sklearn.tree import DecisionTreeClassifier, export_text

# Build a decision tree classifier Set (random_state=42)
...

# Predict on the test set
y_pred = tree_clf.predict(X_test)

# Calculate accuracy
...
print(f"Decision Tree Accuracy: {accuracy:.2f}")

# Display the tree structure
...
print(tree_structure)

```

```python
grader.check("q3")
```


#### 4. Implement Gini Index

Let's manually compute the Gini index for one of the splits in the decision tree.

Here is a description of Gini index from the [lecture notes](https://tools4ds.github.io/DS701-Course-Notes/14-Classification-I-Decision-Trees.html#impurity-measures)


```python
# Function to calculate Gini Index
def gini_index(groups, classes):
    # count all samples at split point
    n_instances = float(sum([len(group) for group in groups]))
    # sum weighted Gini index for each group
    gini = 0.0
    for group in groups:
        ...
    return gini

# Example: Manually splitting data based on feature 0 (sepal length) with threshold 5.5
def test_split(index, value, dataset):
    left, right = [], []
    for row in dataset:
        if row[index] < value:
            left.append(row)
        else:
            right.append(row)
    return left, right

# Combine X and y into a dataset
dataset = np.column_stack((X_train, y_train))

# Test split
left_group, right_group = test_split(0, 5.5, dataset)

# Calculate Gini index for the split
gini = gini_index([left_group, right_group], np.unique(y_train))
print(f"Gini Index for the split: {gini:.3f}")

```

```python
grader.check("q4")
```

#### 5. Prevent Overfitting.

Overfitting happens when the model learns the training data too well, including the noise, which results in poor performance on unseen data. Random Forests can also overfit, especially if the trees are too deep or there are too many trees.

To prevent overfitting, we can:
- Limit the depth of each tree (`max_depth`).
- Reduce the number of trees (`n_estimators`).
- Increase the minimum number of samples required to split a node (`min_samples_split`).
- Increase the minimum number of samples required to be at a leaf node (`min_samples_leaf`).
- Use cross-validation to monitor the model's performance.

Tweak the parameters of the Random Forest to prevent overfitting. The accuracy should be less than 95% & above 80%. Make sure you use `random_state`

```python
# Place the right one
# Set (random_state=42)
...
rf_clf_tuned.fit(X_train, y_train)


# Predict and calculate accuracy for the tuned model
y_pred_rf_tuned = rf_clf_tuned.predict(X_test)
rf_accuracy_tuned = accuracy_score(y_test, y_pred_rf_tuned)
print(f"Tuned Random Forest Accuracy: {rf_accuracy_tuned:.2f}")

```

```python
grader.check("q5")
```

<!-- BEGIN QUESTION -->

#### 6. Interpretability with Decision Trees 

The decision tree is easy to interpret. Write code through the export_graphviz that can create an image to explain the decision tree.

You can refer to the given image provided.


```python
import matplotlib.pyplot as plt
import matplotlib.image as mpimg
# Read the image file and display it using matplotlib
img = mpimg.imread("decision_tree.png")
plt.figure(figsize=(8, 17))
plt.imshow(img)
plt.axis('off')
plt.show()
```

```python
from sklearn.tree import export_graphviz
import graphviz
import matplotlib.pyplot as plt
import matplotlib.image as mpimg


# Use Tree CLF to export the tree structure to a dot file
...

# Display the tree using graphviz
graph = graphviz.Source(dot_data)

# Save the graph to a file
graph.format = 'png'
graph.render("figs/decision_tree")

# Read the image file and display it using matplotlib
img = mpimg.imread("figs/decision_tree.png")
plt.figure(figsize=(8, 17))
plt.imshow(img)
plt.axis('off')
plt.show()
```

```python
grader.check("q6")
```

<!-- END QUESTION -->

#### 7. Interpretability with kNN? 

It can be somewhat difficult to __understand__ why $k$-NN is making a specific prediction. It depends on the data in the neighborhood of the test point.

Let us build a K-NN. Do the following:
- Scale the features
- Build the k-NN with 5 neighbors
- Print the confusion matrix

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
from sklearn.datasets import load_iris
from sklearn.model_selection import train_test_split
from sklearn.preprocessing import StandardScaler
from sklearn.neighbors import KNeighborsClassifier
from sklearn.metrics import accuracy_score, ConfusionMatrixDisplay

# Standardize the features, X_train_scaled, and X_test_scaled
...


# Create the KNN model Set (random_state=42)
...

# Train the model
knn.fit(X_train_scaled, y_train)

# Make predictions
y_pred = knn.predict(X_test_scaled)

# Calculate accuracy
accuracy = accuracy_score(y_test, y_pred)

# Plot the confusion matrix
...
plt.title(f'Confusion Matrix for k-NN (Accuracy: {accuracy:.2f})')
plt.show()

accuracy

```

```python
grader.check("q7")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q3 (points: 2)
```
>>> assert isinstance(tree_clf, DecisionTreeClassifier), 'tree_clf is not an instance of DecisionTreeClassifier'
>>> assert hasattr(tree_clf, 'tree_'), 'DecisionTreeClassifier has not been fitted yet'
>>> assert accuracy > 0.9, 'The accuracy of the Decision Tree Classifier is too low'
```


Tests for q4 (points: 1)
```
>>> (left_group, right_group) = test_split(0, 6.5, dataset)
>>> gini = gini_index([left_group, right_group], np.unique(y_train))
>>> assert type(gini) == float
>>> assert np.round(gini, 3) == 0.584
```


Tests for q5 (points: 1)
```
>>> output = rf_clf_tuned.fit(X_train, y_train)
>>> y_pred_rf = rf_clf_tuned.predict(X_test)
>>> rf_accuracy = accuracy_score(y_test, y_pred_rf)
>>> assert type(rf_accuracy) == float
>>> assert np.round(rf_accuracy, 2) <= 0.95
>>> assert np.round(rf_accuracy, 2) >= 0.8
```


Tests for q6 (points: 1)
```
>>> import os
>>> files_in_figs = os.listdir('figs')
>>> assert 'decision_tree.png' in files_in_figs, "'decision_tree.png' is not in the 'figs' folder."
```


Tests for q7 (points: 2)
```
>>> assert knn.n_neighbors == 5
>>> assert X_train_scaled.shape == X_train.shape, 'X_train_scaled has different shape than X_train.'
>>> assert X_test_scaled.shape == X_test.shape, 'X_test_scaled has different shape than X_test.'
>>> assert X_train is not X_train_scaled, 'X_train and X_train_scaled refer to the same object.'
>>> assert not np.array_equal(X_train, X_train_scaled), 'X_train and X_train_scaled have the same content.'
>>> assert X_test is not X_test_scaled, 'X_test and X_test_scaled refer to the same object.'
>>> assert not np.array_equal(X_test_scaled, X_test), 'X_train and X_train_scaled have the same content.'
```
